# A2-Ciphers

This assignment covers Collections, Loops & Encodings. 

This assignment is out of 8 points, worth 8% of your grade.

## Part 1: Collection Types

This part covers collection types: lists, tuples and dictionaries. 

### QN

Below are the results of a hypothetical experiment of measuring the height of a class in a land far away:

- Mario is 5.6 ft. tall
- Sarai is 5.4 ft. tall
- Demi  is 6.2 ft. tall
- Ian   is 5.8 ft. tall
- Dawn  is 5.7 ft. tall

Using the two provided variable below create a list of strings for the Class name and a list of Floats for their heights.

Remember that the order of lists matters, and so make sure to keep the order of items the same as listed above!

<!-- grade_id: cell-653ac226fd7e65c4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class_names = ["Mario", "Sarai", "Demi", "Ian", "Dawn"]
class_heights = [5.6, 5.4, 6.2, 5.8, 5.7]
```

<!-- grade_id: cell-7ed89774a93617f3, points: 1, grade: true, solution: false -->
```python
# Tests for QN

# Check both lists are defined
assert isinstance(class_names, list)
assert isinstance(class_heights, list)

# Check both lists have the same length, and are of length 5
assert len(class_names) == len(class_heights) == 5

```

Now that we have the data stored in lists, we will explore other collection types. 

## Q2

Create a tuple, that contains just the names of the people in the experiment. Call this tuple 'names_tuple'. 

<!-- grade_id: cell-fb210f0469973e5a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
names_tuple = tuple(class_names)
```

<!-- grade_id: cell-5969238aefebc1dd, points: 1, grade: true, solution: false -->
```python
assert isinstance(names_tuple, tuple)

```

## QN

Next, store the data into a dictionary, where each student's name is the key and their height is the value.

Call this dictionary `results_dictionary`. 

<!-- grade_id: cell-d1afd11046ada78c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
results_dictionary = dict(zip(class_names,class_heights))
```

<!-- grade_id: cell-a823da59fce8614b, points: 1, grade: true, solution: false -->
```python

assert isinstance(results_dictionary, dict)

```

## QN - Lists of Lists

You can also make lists, that are filled with lists! List-ception. 

First, create three different lists:
- A list called `string_list` that containts three strings (can be any strings)
- A list called `number_list` that containts three numbers (can be any numbers - int or float)
- A list called `boolean_list` that containts three boolean (can be any booleans)

Then, create a new list, called `nested_list` which contains the three lists you created above. 

<!-- grade_id: cell-49a40611dcdbbea0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
string_list = ['a', 'b', 'c']
number_list = [1, 2, 3]
boolean_list = [True, True, True]
nested_list = [string_list, number_list, boolean_list]
```

<!-- grade_id: cell-51ffc8b55e7f36fc, points: 1, grade: true, solution: false -->
```python
assert isinstance(string_list, list)
assert len(string_list) == 3
assert isinstance(string_list[0], str)

assert isinstance(number_list, list)
assert len(number_list) == 3
assert isinstance(number_list[0], int) or isinstance(number_list[0], float)

assert isinstance(boolean_list, list)
assert len(boolean_list) == 3
assert isinstance(boolean_list[0], bool)
```

<!-- grade_id: cell-04845071d9603b3d, points: 1, grade: true, solution: false -->
```python
assert isinstance(nested_list, list)
assert len(nested_list) == 3
assert isinstance(nested_list[0], list)
```

## Part 2: Working With Collections

This part covers working with collections - indexing and 

```python
# For QN &QN, the following lists are provided for you
list_1 = [10, 20, 40, 50]
list_2 = [13, 15, 17, 19, 22, 25]
list_3 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
list_4 = [21, 9, 98, 289, 938]
```

## Qn - Indexing

Do the following using Indexing:
- Store the first value in list_1 to index_1
- Store the first four values of list_2 to index 2
- Store the last four values of list_3 to index 3 

<!-- grade_id: cell-650e61f820451347, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
index_1_ans = list_1[0]
index_2_ans = list_2[0:4]
index_3_ans = list_3[-4:]

```

<!-- grade_id: cell-2468c24a17df0b96, points: 1, grade: true, solution: false -->
```python
assert isinstance(index_1_ans, int)

```

<!-- grade_id: cell-3ae0d083dd79537d, points: 1, grade: true, solution: false -->
```python
assert isinstance(index_2_ans, list)
assert len(index_2_ans) == 4

```

<!-- grade_id: cell-f591dd9f036b0ffe, points: 1, grade: true, solution: false -->
```python
assert isinstance(index_3_ans, list)
assert len(index_3_ans) == 4

```

## QN - Comparisons using Indexing

Do the following Comparisons using Indexing:
- Check if the first value of list_1 is one of the last three values in list_3 (Hint: remember operator 'in')
    - Store the value in comp_result_1
- Check if the third value of list_4 is greator than the fourth value in list_4
    - Store the value in comp_result_2
- Check if the last value of list_4 is less than the second value in list_4
    - Store the value in comp_result_3
- Check if the second value of list_2 multiplied by the seventh value of list_3 is greator than the second to last value in list_4
    - Store the value in comp_result_4

<!-- grade_id: cell-58d83066fbe5dfa8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
comp_result_1 = index_1_ans in list_3[-3:]
comp_result_2 = list_4[2] > list_4[3]
comp_result_3 = list_4[len(list_4)-1] < list_4[1]
comp_result_4 = (list_2[1] * list_3[6]) > list_4[-1]
```

<!-- grade_id: cell-2209a4edc72a8364, points: 1, grade: true, solution: false -->
```python
# Tests for QN
assert isinstance(comp_result_1, bool)

```

<!-- grade_id: cell-6353d9279a80e98f, points: 1, grade: true, solution: false -->
```python
assert isinstance(comp_result_2, bool)

```

<!-- grade_id: cell-d75011cb18874b17, points: 1, grade: true, solution: false -->
```python
assert isinstance(comp_result_3, bool)

```

<!-- grade_id: cell-244b969301ac73e8, points: 1, grade: true, solution: false -->
```python
assert isinstance(comp_result_4, bool)

```

## Part 3: Loops

This part covers using loops. 

## QN

For each of the scenarios below, indicate whether it is a procedure best done with a for loop, a while loop, or if the task is not something best done with a loop. 

- a) You want to apply a transform to every piece of data in a list. 
- b) You want to continuously update the position of a character displayed on a screen, as long as the status of the character is 'alive'.
- c) You want to check three different value comparisons are all true together, each checking a different variable. 
- d) You want your computer to connect to an external network. As long as your connection status is False, you want to try and re-connect. 

<!-- grade_id: cell-eb8a52c2ab3e909b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
loop_q_a = "for"
loop_q_b = "while"
loop_q_c = "not_a_loop"
loop_q_d = "while"
```

<!-- grade_id: cell-78ffc841fc725614, points: 1, grade: true, solution: false -->
```python
assert loop_q_a in ['for', 'while', 'not_a_loop']

```

<!-- grade_id: cell-b4743cbf55530c66, points: 1, grade: true, solution: false -->
```python
assert loop_q_b in ['for', 'while', 'not_a_loop']

```

<!-- grade_id: cell-474f765b4c8b51bb, points: 1, grade: true, solution: false -->
```python
assert loop_q_c in ['for', 'while', 'not_a_loop']

```

<!-- grade_id: cell-3f8ff16491f2e035, points: 1, grade: true, solution: false -->
```python
assert loop_q_d in ['for', 'while', 'not_a_loop']

```

# QN

What is the sum of all even numbers between 19 and 35?

Answer this question using a for loop, and 

Hints:



<!-- grade_id: cell-749ae260cbdc082d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
int_list = list(range(19,36))
sum = 0
for num in int_list:
    sum = sum + num
print(sum)
```

Output:
```
459
```

<!-- grade_id: cell-6b415de65a9332be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

## QN

Write a while loop that adds the first 100 numbers (from 1 to 100) together.

To do so, use a while loop that continues while the condition of `cur_num` being less than or equal to `stop_num`. 

Inside the loop, it should at the `cur_num` to total, and also add 1 to `cur_num`. 

<!-- grade_id: cell-fdabc8d0323473dd, points: 0, grade: false, solution: true -->
```python
# These variables provided for you - use them in your solution
total = 0
cur_num = 1
stop_num = 100

while cur_num < 100:
    total = total+cur_num
    cur_num = cur_num+1

```

<!-- grade_id: cell-f6513546d3afb423, points: 0, grade: true, solution: false -->
```python
assert isinstance(cur_num, int)

```

## Example Provided

In the following examples we will need to use an operation to add something to a list. Some examples for doing so are provided. 

Given a list, `my_list` to add an item `my_item` to the end of the list, you can do `my_list.append(my_item)`. 

```python
# Example of appending to a list
my_list = ['a', 12, None]
print('Before Append: ', my_list)
my_list.append(True)
print('After  Append: ', my_list)
```

```python
# Example of using append with loops
list_of_values = [0, 1, 2, 3]

# Initialize an empty list that will be used to store outputs
outputs = []

# Loop through all items in `list_of_items`
for item in list_of_values:
    
    # Add incremented item to the list of outputs
    outputs.append(item + 1)

# Check what outputs ends up as
print(outputs)
```

## QN

Using `num_list`, write a for loop to loop across each value, multiplying it by -1. 

Store the result for each value, in a list called `eval_1`, by defining `eval_1` before the loop, and using `append` inside the loop. 

```python
# This creates a list of all numbers from 1-12, inclusive
num_list = list(range(1, 13))
```

<!-- grade_id: cell-c3c8595879347f24, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
eval_1 = []
for num1 in num_list:
    eval_1.append(num1 * (-1))
```

<!-- grade_id: cell-09e1d8d99d175d45, points: 1, grade: true, solution: false -->
```python
assert isinstance(eval_1, list)
assert isinstance(eval_1[0], int)

```

## Part 4: Pulling it Together - Adding Conditionals

This part covers integrating collections and loops with conditional and operators. 

## QN

Write some code that uses a loop to find all the values from `my_list` that are perfectly divisible by 3. 

Use a for loop to do this, and save any values from `my_list` that meet this condition into a new list called `div_3`. 

```python
# This list provided to use for the following questions
data_list = [ -1, 20, -100, -40, 33, 97, -101, 45, -79, 96]
```

<!-- grade_id: cell-b3adf3ba9e29eb9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
div_3 = []
for num2 in data_list:
    if num2 % 3 == 0:
        div_3.append(num2)
```

<!-- grade_id: cell-bb9b4796b7db8253, points: 1, grade: true, solution: false -->
```python
# Tests for QN
assert isinstance(div_3, list)
assert(len(div_3) == 3)

```

## QN

Write a for loop that is going to check whether the values in `data_list` are even or odd. 

Use a for loop to do this. For each value, it should add `True` to a new list `is_even` if the value is even, and `False` to `is_even` if the value is odd. 

<!-- grade_id: cell-af4c9eabb79396f1, points: 0, grade: false, solution: true -->
```python
is_even = []
for num3 in data_list:
    if num3 % 2 == 0:
        is_even.append(True)
    else:
        is_even.append(False)
```

<!-- grade_id: cell-3b8896dd294ccc5a, points: 1, grade: true, solution: false -->
```python
# Tests for QN
assert isinstance(is_even, list)
assert(len(is_even) == len(data_list))

```

## QN

Write a for loop that will check each value list_for_1 to see if they are even or odd and whether they are positive or negative.
- If they are both positive and even, encode this as the number 1
- If they are both negative and odd, encode this as the the number -1
- Otherwise, encode this as the number 0

- Store the result for each value into a list called `num_type`

<!-- grade_id: cell-c55e1457c131cb10, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

num_type =[]

for num4 in data_list:
    if (num4 % 2 == 0) and (num4 > 0):
        num_type.append(1)
        print('if1')
    if(num4 % 2 != 0) and (num4 < 0):
        num_type.append(-1)
        print('if2')
    else:
        num_type.append(0)
        print('else')
        
print(num_type)
```

Output:
```
if2
if1
else
else
else
else
else
if2
else
if2
if1
else
[-1, 1, 0, 0, 0, 0, 0, -1, 0, -1, 1, 0]
```

<!-- grade_id: cell-552f3b912ed4e050, points: 1, grade: true, solution: false -->
```python
# Tests for QN
assert isinstance(num_type, list)
assert(len(num_type) == len(data_list))

```

### Looping through dictionaries

```python
# This example of looping through dictionaries provided
dictionary = {
    'key1' : 'value1', 
    'key2' : 'value2'
}

for item in dictionary:
    
    # When you loop through a dictionary like this, `item` is the key 
    key = item
    
    # Using the key, you can access the associaed value of the dictionary from within the loop
    value = dictionary[item]
    
    # Print out data for each loop iteration
    print('New Loop Iteration')
    print('\tKey is:\t', item)
    print('\tVal is:\t', value)
```

Output:
```
New Loop Iteration
	Key is:	 key1
	Val is:	 value1
New Loop Iteration
	Key is:	 key2
	Val is:	 value2
```

# QN

Using the dictionary `subjects` write a for loop that loops across the dictionary and collects all subject numbers (ex. 'S1') where the dictionary value is False. 

Imagine, for example, the dictionary indicates whether processing is complete, and we wanted to get a list of the subjects that need more analyses. 

To answer this question, use a for loop across the `subjects` dictionary. You then need to get the associated value in each iteration, and check if it is `True`. If it is `True`, you can use `continue` to skip ahead to the next iteration. Otherwise, append the subject number (ex 'S1') to a list called `not_finished`. 

```python
# This dictionary provided for you
subjects = {
    'S1' : True,
    'S2' : False,
    'S3' : True,
    'S4' : False,
}
```

<!-- grade_id: cell-1118d6a4ac43e87d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
not_finished = []
for item1 in subjects:
    key = item1
    value = subjects[item1]
    if value == True:
        continue
    else:
        not_finished.append(key)
```

<!-- grade_id: cell-36d719f579df1fe3, points: 1, grade: true, solution: false -->
```python
assert isinstance(not_finished, list)
assert len(not_finished) == 2

```

## Part 5: Application - Ciphers

This part is an application of all the tools we have covered so far, to the problem of creating and using ciphers. 

A cipher is procedure for encoding and decoding secret messages (or encrypted data). 

To work on this application, we are going to need to use loops and collections. 

It also uses concepts and ideas from encoding and decoding - and in particular the operators `ord` and `chr`. 

```python
# Setup - run this cell before doing the next part of the assignment
import os
if not os.path.exists('A2Code'):
    os.mkdir('A2Code')
```

```python
# Example / outline of writing a cipher
#   The following shows a layour of writing a cipher.
#   You will be writing code to actually do the encoding / decoding.

# Create a message
original_message = 'Hello from Tom.'
print('Original Message:\t', original_message)

# First, Apply Encoding Code, to convert to an encoded message
encoded = 'ĐĭĴĴķèĮĺķĵèĜķĵö'
print('Encoded Message:\t', encoded)

# Then, Apply Decoding Code, to decode into an interpretable message
decoded = 'Hello from Tom.'
print('Decoded Message:\t', decoded)
```

Output:
```
Original Message:	 Hello from Tom.
Encoded Message:	 ĐĭĴĴķèĮĺķĵèĜķĵö
Decoded Message:	 Hello from Tom.
```

## QN - Creating an Encoder

Write a code block to encode strings into a secret message.

To do so:
- Initialize a variable called `encoded` as an empty string
- Use a for loop to loop across all characters of a presumed string variable called `message`
    - Inside the loop, convert the character to another character:
        - Get the unicode code point for the character (using `ord`)
        - Add the value of `key` to that code point (which should be an int)
        - Convert this new int back to a character (using `chr`).
    - Also inside the loop, add this converted char to the string `encoded`

<!-- grade_id: cell-d77bc7ed64441ebc, points: 0, grade: false, solution: true -->
```python
%%writefile A2Code/encoder.py

# YOUR CODE HERE
length = len(message)
index = 0
encoded = ''
while index < length:
    letter = ord(message[index])
    encLetter = letter + key
    newLetter = chr(encLetter)
    encoded = encoded + newLetter
    index = index + 1

```

Output:
```
Overwriting A2Code/encoder.py
```

<!-- grade_id: cell-8dbc526034a66ab8, points: 1, grade: true, solution: false -->
```python
key = 200
message = 'hello'

%run -i ./A2Code/encoder.py  

assert isinstance(encoded, str)
assert encoded == 'İĭĴĴķ'

print('Original Message: \t', message)
print('Encoded Message: \t', encoded)
```

Output:
```
Original Message: 	 hello
Encoded Message: 	 İĭĴĴķ
```

## QN - Decoder

Write a code block to decode strings from secret encodings back to readable messages.

To do so:
- Initialize a variable called `decoded` as an empty string
- Use a for loop to loop across all characters of a presumed string variable called `encoded`
    - Inside the loop, convert the character to another character:
        - Get the unicode code point for the character (using `ord`)
        - Subtract the value of `key` to that code point (which should be an int)
            - This undoes the secret encoding of the character, getting back to the original value
        - Convert this new int back to a character (using `chr`).
    - Also inside the loop, add this converted char to the string `decoded`
    
Note that the code to answer this question should look very similar to how you answered the encoding question. 

In fact, you can even copy over the encoding code, and make some small changes to make it function as a decoder. 

<!-- grade_id: cell-b629dac4beccd5f1, points: 0, grade: false, solution: true -->
```python
%%writefile A2Code/decoder.py

# YOUR CODE HERE
decoded = ''
length = len(encoded)
index = 0
while index < length:
    letter = ord(encoded[index])
    decLetter = letter - key
    newLetter = chr(decLetter)
    decoded = decoded + newLetter
    index = index + 1

```

Output:
```
Overwriting A2Code/decoder.py
```

<!-- grade_id: cell-ddf1c25663771ad8, points: 1, grade: true, solution: false -->
```python
key = 200
encoded = 'İĭĴĴķ'

%run -i ./A2Code/decoder.py  

assert isinstance(decoded, str)
assert decoded == 'hello'

print('Encoded Message: \t', encoded)
print('Decoded Message: \t', decoded)
```

Output:
```
Encoded Message: 	 İĭĴĴķ
Decoded Message: 	 hello
```

### Using the encoder and decoder together

The following are a couple examples of using the encoder and decoder together. 

```python
# Using the encoder
key = 500
message = 'This is a secret message'

%run -i ./A2Code/encoder.py  
%run -i ./A2Code/decoder.py

print('\nOriginal Message: \t', message)
print('\nEncoded Message: \t', encoded)
print('\nDecoded Message: \t', decoded, '\n')

assert message == decoded
```

Output:
```
Original Message: 	 This is a secret message

Encoded Message: 	 ɈɜɝɧȔɝɧȔɕȔɧəɗɦəɨȔɡəɧɧɕɛə

Decoded Message: 	 This is a secret message
```

```python
# Using the encoder
key = 1000
message = 'Python is pretty cool, eh?'

%run -i ./A2Code/encoder.py  
%run -i ./A2Code/decoder.py

print('\nOriginal Message: \t', message)
print('\nEncoded Message: \t', encoded)
print('\nDecoded Message: \t', decoded, '\n')

assert message == decoded
```

Output:
```
Original Message: 	 Python is pretty cool, eh?

Encoded Message: 	 иѡќѐїіЈёћЈјњэќќѡЈыїїєДЈэѐЧ

Decoded Message: 	 Python is pretty cool, eh?
```

## Using a Custom Encoder

The encoder & decoding we did above is built on conversions based on manipulating uncide code points.

If a potential attacker wanted to decode our messages, and happened to know we were using this approach, there are some relatively simple ways that they could try to figure out our encoding key, and then, if successful decode our messages. 

Thought experiment: can you think of ways to decode messages that use this encoding?

To thwart this potential attacker, let's now explore using a custom encoding. Instead of manipulating code points, we are going to use a custom mapping to switch out characters to encode our messages. 

```python
# Custom encoding dictionary provided - make sure you run this cell
#   In thish encoding, each key will get switched out for the associated value
#     So, for example, whenever we see an 'a', we want to switch it out for an 'r'. 
custom_encodings = {
    'a' : 'r', 
    'e' : 'p', 
    'i' : 'm',
    'o' : 'n',
    'u' : 's'
}
```

## QN

Write a code block to encode strings from original messages, using our custom encoding.

To do so:
- Initialize a variable called `custom_encoded` as an empty string
- Use a for loop to loop across all characters of a presumed string variable called `custom_message`
    - Inside the loop, check if the current character is in the `custom_encodings` dictionary
        - If it is, use the current char to get the value from `custom_encodings`, and concatenate it to `custom_encoded`
            - The line inside the if should look something like `out = out + dictionary[char]
        - Otherwise, concatenate the current character to `custom_encoded`

In fact, you can even copy over the encoding code, and make some small changes to make it function as a decoder. 


<!-- grade_id: cell-c71c52f6393468cd, points: 0, grade: false, solution: true -->
```python
%%writefile A2Code/custom_encoder.py

# YOUR CODE HERE
custom_encoded = ''
for char in custom_message:
    if char in custom_encodings:
        custom_encoded = custom_encoded + (custom_encodings[char])
    else:
        custom_encoded = custom_encoded +char
```

Output:
```
Overwriting A2Code/custom_encoder.py
```

<!-- grade_id: cell-66b937622cc2a5f6, points: 1, grade: true, solution: false -->
```python
custom_message = 'vowels aint great'

%run -i ./A2Code/custom_encoder.py

assert isinstance(custom_encoded, str)
assert custom_encoded == 'vnwpls rmnt grprt'

print('\nOriginal Message: \t', custom_message)
print('\nEncoded Message: \t', custom_encoded, '\n')
```

Output:
```
Original Message: 	 vowels aint great

Encoded Message: 	 vnwpls rmnt grprt
```

```python
## QN - Encodings with variable keys
```

<!-- grade_id: cell-06e54039038018ff, points: 0, grade: false, solution: true -->
```python
%%writefile A2Code/variable_encoder.py

length = len(message)
index = 0
encoded = ''
key=start_key
while index < length:
    letter = ord(message[index])
    encLetter = letter + key
    newLetter = chr(encLetter)
    encoded = encoded + newLetter
    index = index + 1
    key = key + key_increment

```

Output:
```
Overwriting A2Code/variable_encoder.py
```

<!-- grade_id: cell-929fe5d77dbe70be, points: 1, grade: true, solution: false -->
```python
message = "By the way, what we're doing here basically the same as encryption"

start_key = 150
key_increment = 3

%run -i ./A2Code/variable_encoder.py

assert isinstance(encoded, str)
assert encoded == "ØĒ¼ēĊĊÈĢďĪà×ıĥġķæŀıöńĺøĿōŊŒŎĊŕŕťśęŞŠŵŮūŬźŽƍķƎƅƅŃƙƊƙƔŒƖƫśƣƯƧƹǃƽǄƼǅǇ"

print('\nOriginal Message: \t', message)
print('\nEncoded Message: \t', encoded, '\n')
```

Output:
```
Original Message: 	 By the way, what we're doing here basically the same as encryption

Encoded Message: 	 ØĒ¼ēĊĊÈĢďĪà×ıĥġķæŀıöńĺøĿōŊŒŎĊŕŕťśęŞŠŵŮūŬźŽƍķƎƅƅŃƙƊƙƔŒƖƫśƣƯƧƹǃƽǄƼǅǇ
```

## Breaking a Variable Encoder

### Extra (optional) Information for the Curious - on Cryptography & Alan Turing

This kind of 'variable encoder' in which the encoding transform updates every time it is applied is much harder for an outside attacker to break (if they do not know the initial key, and the way it changes every application). 

This is, broadly, the kind of encryption which was used by German forces in World War II to encode messages - the ENIGMA code. 

This approach is, however, still breakable - it just takes a lot of clever analysis and number crunching. During World War II, Alan Turing, having already at that point done fundamental work developing the very concept of a digital computer, developed machines and algorithms to try to break the German Enigma code, eventually managing to decode it with his team at Bletchley Park.

Cracking this code was a huge development in the war - it is estimated this code breaking work, which allowed the Allies to decode messages and anticipate German actions, may have shortened the war by up 2 to 3 years, saving millions of lives in the process.

By the early 1950s, though his codebreaking work was still stricly confidential and he was largely unknown by the general public, Alan Turing had gone on to make other huge contributions to the fields of computation and artificial intelligence, amongst others, that would go on to have massive impact on the modern world. 

The British government, who years before had employed him as a codebreaker, then prosecuted Alan Turing for homosexual activity, and forced him to undergo chemical castration. Alan Turing died of suicde in 1954, at the age of 41. 

For more information on Alan Turing, check out the article linked below, and/or look into the movie `The Imitation Game`, a dramatization of his work on code breaking during the war. 

Short Article on Alan Turing: https://www.bbc.com/news/technology-18419691

## Variable Decoder

The code below is provided for you, that performs decoding of a variable decoder. 

To do so, it needs to know the `start_key` and `key_increment` that were used.

```python
%%writefile A2Code/variable_decoder.py

# Knowing the start key, and the length of the message, we can reconstruct the current key
key = start_key + (len(encoded) * key_increment)

decoded = ''

# Note that this decodes the message backwards, stepping back along key values
for char in encoded[::-1]:
    
    # Step the key back one step, and apply to current character
    key = key - key_increment
    decoded = decoded + chr(ord(char) - key)
    
# Having decoded backwords, flip the message back around
decoded = decoded[::-1]
```

Output:
```
Overwriting A2Code/variable_decoder.py
```

## Decoding a provided message

The cell below has an encoded message, using variable encoding, and the `start_key` and `key_increment` needed to decode. 

Run the following cell to execute the decoding code from above, and decode the message. 

```python
# Do decoding on message above
encoded = """Ãâø¡ü¬èõ÷­¼±ÜµÿĈċĂ¿ĚĒĚÇĊĝĒÏĖġğĦĲĤīĦáķĭĬéĬŀłĺĺŃńľŉőĀëíĪŐŝœŒšđŪŖŰęŴŬŴġŤŷŬĩůżŽŶĳƃƆƐŇĽƒƐŃƈƖ\
ƗƒƟƐƥƦŕƦƧśƱƧƢƷŦőœƍǆůǅƻƺŷǐƼǖƋƁǌǋƇǢǚǢƏǙǔǥǧǞǩƝǳǰƣǷǬǪǷǹȈƱǿǾȂǾƻȑȇȊȖǅȊȒțȕȔȣǓȨȫȮȡȣǟǮǣȫȶȵȷȼɆȺɁȼǷɎɋǽɎɏȃəɏɒ\
ɞȍɢɥɨɛɝșɲɦɫɭȣɧɬȩɬȭɿʃʂɿɼɼʏȽʎʑʗʎʖʗəȷȹɷʥʤʤɥɛʑʮʮ"""

start_key = 123
key_increment = 2

%run -i ./A2Code/variable_decoder.py

print(decoded)
```

Output:
```
Hey y'all - I hope you are enjoying the assignment!

Either way you are done now, so congrats on that!

By the way, if you happen to really like this cipher stuff - following up on this stuff will be a project option.

From, Tom
```

## The End!

This is the end of the assignment!
Have a look back over your answers, and also make sure to Restart & Run All from the kernel menu to double check that everything is working properly. 

When you are ready to submit your assignment, upload it to TritonED under Assignment-2.